```python
import numpy as np
from helpers import print_table
```

<!-- grade_id: cell-1486738b77c0d868, points: 0, grade: false, solution: true -->
```python
# copy inference_by_enumeration from Problem 1 & print and compare the probability tables here!

# YOUR CODE HERE
def inference_by_enumeration(FJDT: np.ndarray, query_variable_indices: tuple, evidence_variable_indices: tuple=tuple()) -> np.ndarray:
    '''
    Computes the answer to a probabilistic query exactly from the full joint distribution table.
    :param table: The full joint distribution table as a np.ndarray.
    :param query_variable_indices: A tuple containing the indices of the query variables in the FJDT.
    :param evidence_variable_indices: A tuple containing the indices of the evidence variables in the FJDT.
    :returns: The answer to the probabilistic query; a `np.ndarray`.
    ''' 
    assert type(FJDT) == np.ndarray, "FJDT must be a np.ndarray"
    assert type(query_variable_indices) == tuple, "query_variable_indices must be a tuple"
    assert type(evidence_variable_indices) == tuple, "evidence_variable_indices must be a tuple"
    
    # compute the set of non-query and non-evidence variables, Z
    query_variables = query_variable_indices + evidence_variable_indices
    Z = tuple(set(range(FJDT.ndim)).difference(query_variables))

    # YOUR CODE HERE
    p1 = np.sum(FJDT, axis=Z, keepdims=True)
    normConst = np.sum(p1, axis=query_variable_indices, keepdims=True)
    return p1 / normConst
```

# 2. Pairwise vs. Mutual Independence (4 points)

**Definition**: We say that two random variables are *pairwise independent* if $$p(X_n \mid X_m) = P(X_n)$$ and hence $$p(X_m, X_n) =  p(X_n \mid X_m) p(X_m) = p(X_n) p(X_m) $$

**Definition**: We say that $n$ random variables are *mutually independent* if $$p(X_i \mid X_{S}) = p(X_i)\;\; \forall S \subseteq \{1, \dots, n\} \setminus \{ i \}$$ and hence $$p(X_{1:n}) = \prod_{i=1}^n p(X_i)$$




<div class="alert alert-warning">
Show that pairwise independence between all pairs of variables does not necessarily imply mutual independence. Come up with a minimal counterexample that has exactly three binary random variables. (4 points)
</div>

Specify this counterexample via its full joint distribution table (FJDT). **Briefly** outline your thought process in the text field below (use $\LaTeX$ and markdown) and store the model's full joint distribution table into the `XYZ` variable. It is sufficient to show pairwise independence and non-mutual independence by comparing products of marginals and joint distributions. 

**Hint**: Copy your implementation of `inference_by_enumeration` from Problem 1. You can use `print_table` to visualize your distribution tables such as the FJDT, products of marginals, and joint distributions.

```python
help(print_table)
```

Output:
```
Help on function print_table in module helpers:

print_table(probability_table: numpy.ndarray, variable_names: str) -> None
    Prints a probability distribution table.
    
    Parameters
    ----------
    probability_table : np.ndarray
        The probability distribution table
    variable_names : str
        A string containing the variable names, e.g., 'CDE'.
    
    Returns
    -------
    None
```

<!-- grade_id: cell-6ced5d0ec1cbd764, points: 1, grade: true, solution: true -->
When thinking of the three variables X,Y,Z, we could assign them each a unit square inside a $[0,1] \times [0,1]$ area. Therefore, we have a probability for A, B, C of
$$P(A) = P(B) = P(C) = \frac{1}{2}$$

Furthermore, it should be noted that:
$$P(A \cap B) = P(A \cap C) = P(B \cap C) = \frac{1}{2}$$

From this the following can be said:
$$P(A \cap B) = P(A)*P(B) \\
P(A \cap C) = P(A) * P(C) \\
P(B \cap C) = P(B) * P(C)  $$

However, there is a contradiction forming, because $P(A \cap B \cap C) = P(A \cap B = \frac{1}{4})$, yet $P(A)*P(B)*P(C) = \frac{1}{8}$. From this can be concluded that A, B and C are not independent from each other, as this would have resultet in the same values for both $P(A \cap B \cap C)$ and $P(A)*P(B)*P(C)$

<!-- grade_id: cell-7f0b4fa0aa65bb39, points: 0, grade: false, solution: true -->
```python
XYZ = np.ones((2,2,2))

# YOUR CODE HERE
XYZ = np.zeros((2,2,2))
X, Y, Z = 0, 1, 2

XYZ[0,0,0] = 1/4
XYZ[0,1,1] = 1/4
XYZ[1,0,1] = 1/4
XYZ[1,1,0] = 1/4

print("FJDT")
print_table(XYZ, "XYZ")

print("Single")
print_table(inference_by_enumeration(XYZ, (X,), ()).squeeze(), "X")
print_table(inference_by_enumeration(XYZ, (Y,), ()).squeeze(), "Y")
print_table(inference_by_enumeration(XYZ, (Z,), ()).squeeze(), "Z")

print("Joint")
print_table(inference_by_enumeration(XYZ, (X,Y), ()).squeeze(), "XY")
print_table(inference_by_enumeration(XYZ, (X,Z), ()).squeeze(), "XZ")
print_table(inference_by_enumeration(XYZ, (Y,Z), ()).squeeze(), "YZ")

print("Marginals multiplied")
print_table((inference_by_enumeration(XYZ, (X,), ()) * inference_by_enumeration(XYZ, (Y,), ())).squeeze(),"XY")
print_table((inference_by_enumeration(XYZ, (X,), ())* inference_by_enumeration(XYZ, (Z,), ())).squeeze(), "XZ")
print_table((inference_by_enumeration(XYZ, (Y,), ()) * inference_by_enumeration(XYZ, (Z,), ())).squeeze(), "YZ")

print("P(A, B, C) = FJDT")
print_table(inference_by_enumeration(XYZ, (X,Y,Z,), ()), "XYZ")
print("P(A)*P(B)*P(C)")
a_b_c = inference_by_enumeration(XYZ, (X,), ()) * inference_by_enumeration(XYZ, (Y,), ()) * inference_by_enumeration(XYZ, (Z,), ())
print_table(a_b_c.squeeze(), "XYZ")
```

Output:
```
FJDT
<IPython.core.display.HTML object>
Single
<IPython.core.display.HTML object>
<IPython.core.display.HTML object>
<IPython.core.display.HTML object>
Joint
<IPython.core.display.HTML object>
<IPython.core.display.HTML object>
<IPython.core.display.HTML object>
Marginals multiplied
<IPython.core.display.HTML object>
<IPython.core.display.HTML object>
<IPython.core.display.HTML object>
P(A, B, C) = FJDT
<IPython.core.display.HTML object>
P(A)*P(B)*P(C)
<IPython.core.display.HTML object>
```

<!-- grade_id: cell-b612791dd41278bf, points: 2, grade: true, solution: false -->
```python
assert XYZ.shape == (2, 2, 2), 'FJDT must have shape (2,2,2)'
assert XYZ.sum() == 1, 'Probabilites in FJDT must sum to one'

```

```python

```